## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc, to develop a new version of the function with an improved interface (the interface is the arguments that are passed to the function) with fewer arguments. Add appropriate checks inside your function on the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
def area(x0,y0,x1,y1,x2,y2):
    return abs(0.5*(x0*(y1-y2)+x1*(y2-y0)+x2*(y0-y1)))
point1=(1,1)
point2=(1,2)
point3=(2,1)
x=[]
x.append(point1[0])
x.append(point1[1])
x.append(point2[0])
x.append(point2[1])
x.append(point3[0])
x.append(point3[1])
print(x)
print(area(x[0],x[1],x[2],x[3],x[4],x[5]))
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that, given only the vertex coordinates, computes the area of a simple polygon. Test your function for some simple shapes. Pay careful attention to the range of any loops. **Test for polygons without a vertex at the origin.**

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
xcors=[0,1,0,0]
ycors=[0,0,1,0]

def area(xcors,ycors):
    a=0
    for i in range(len(xcors)-1):
        a+=((xcors[i]*ycors[i+1])-(xcors[i+1]*ycors[i]))/2
    return abs(a)

print(area(xcors,ycors))
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match. Print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    result=[]
    if len(x)!=len(y):
        return "vectors aren't the same size"
    for i in range(len(x)):
        result.append(x[i]+y[i])
    return result

```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
# tests ##
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Optional (advanced) 

Try writing a one-line version that performs the `sum_vector` operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
#idk zip func
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A Python dictionary from college abbreviation to name; and
1. A Python list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
collages={'homerton': "HO", 'Magdalene': 'M', "King's" : "K","Queens":"Q","St Catharine's":"CTH"} 
def abrev(collage):
    return collages[collage]
print(abrev("homerton"))

lis=list(collages)
done=sorted(lis)
print(done)
```

Output:
```
HO
["King's", 'Magdalene', 'Queens', "St Catharine's", 'homerton']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
...
```

```python

```